This assigmaen deals with the particle in a box

$$ \phi_i=\int_0^1{xdx} $$

<!-- grade_id: suma, points: 0, grade: false, solution: true -->
```python
def suma(a,b):
    '''
    escriba la suma
    '''
    return a+b
    raise NotImplementedError()
```

<!-- grade_id: suma2, points: 5, grade: true, solution: false -->
```python
suma(2,4)

assert 6
```

```python

```